```python
# Initialize Otter
import otter
grader = otter.Notebook("example.ipynb")
```

```python
import matplotlib.pyplot as plt
import numpy as np
```

<!-- BEGIN QUESTION -->

**Question 1.** Assign `x` to the smallest prime number.

_Points:_ 16

```python
x = ...
```

```python
grader.check("q1")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.** Visualize the answer

_Points:_ 1.5

```python
## solution ##
plt.plot(...);
```

<!-- END QUESTION -->

This cell is not part of a question.

```python
y = 3
```

**Question 3.** Define `square` and assign `nine` to 3 squared.

_Points:_ 12

```python
def square(x):
    y = ...
    ...

nine = square(3)
```

```python
grader.check("q3")
```

<!-- BEGIN QUESTION -->

**Question 4.** What does equilateral mean?

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# this isn't part of a question
# it's here to make sure that we get a MD cell above to close the export
# of question 4
```

**Question 5.** Approximate the area and circumference of a circle with radius 3.

_Points:_ 1

```python
pi = 3.14
if True:
    ...
    print('A circle with radius', radius, 'has area', area)

def circumference(r):
    # Next, define a circumference function.
    pass
```

```python
# This question has no tests.
```

<!-- BEGIN QUESTION -->

**Question 6.** Write something

_This question has a custom prompt below, so that prompt should be in the output. It also has no solution!_

_Points:_ 100

_Write your thing here._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 7:** What is the answer?

_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 8:** Test intercell seeding by generating 10 random $N(4,2)$ numbers.

_Points:_ 12

```python
z = ...
z
```

```python
grader.check("q8")
```

**You're done!**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(force_save=True)
```

Tests for q1 (points: None)
```
>>> isinstance(x, int)
True
```
failure_message: This is not an int.
success_message: Congrats you passed this test case!\
```
>>> None
```
success_message: Congrats, this passed!
```
>>> 0 < x < 100
True
```
failure_message: This should have passed.
success_message: Congrats your x value is in the correct range!


Tests for q3 (points: None)
```
>>> nine
9
```
failure_message: This fails!
success_message: This works!
```
>>> square(16)
256
```
success_message: Congrats you passed this test case!


Tests for q8 (points: None)
```
>>> len(z) == 10
True
```
failure_message: Make sure the length is 10.
success_message: The length is correct!
