```python
# Initialize Otter
import otter
grader = otter.Notebook("master.ipynb")
```

```python
import matplotlib.pyplot as plt
import numpy as np
```

<!-- BEGIN QUESTION -->

**Question 1.** Assign `x` to the smallest prime number.

```python
x = ...
```

```python
grader.check("q1")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.** Visualize the answer

```python
plt.plot(...);
```

<!-- END QUESTION -->

This cell is not part of a question.

```python
y = 3
```

**Question 3.** Define `square` and assign `nine` to 3 squared.

```python
def square(x):
    y = ...
    ...

nine = square(3)
```

```python
grader.check("q3")
```

<!-- BEGIN QUESTION -->

**Question 4.** What does equilateral mean?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# this isn't part of a question
# it's here to make sure that we get a MD cell above to close the export
# of question 4
```

**Question 5.** Approximate the area and circumference of a circle with radius 3.

```python
pi = 3.14
if True:
    ...
    print('A circle with radius', radius, 'has area', area)

def circumference(r):
    # Next, define a circumference function.
    pass
```

```python
# This question has no tests.
```

<!-- BEGIN QUESTION -->

**Question 6.** Write something

_This question has a custom prompt below, so that prompt should be in the output. It also has no solution!_

_Write your thing here._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 7:** What is the answer?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**You're done!**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> isinstance(x, int)
True
```
failure_message: This is not an int.
success_message: Congrats you passed this test case!\
```
>>> None
```
success_message: Congrats, this passed!
```
>>> 0 < x < 100
True
```
failure_message: This should have passed.
success_message: Congrats your x value is in the correct range!


Tests for q3 (points: None)
```
>>> nine
9
```
failure_message: This fails!
success_message: This works!
```
>>> square(16)
256
```
success_message: Congrats you passed this test case!
